Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "GEORGIA TAKATZOGLOU"
ID = "230"
```

---

### Part A: Regular expressions (3 points)

In the exercised of this part you should NOT assume that gutenberg files are locally available in the file system.

```python
import re
import nltk
nltk.download('gutenberg')
nltk.download('punkt_tab')
from nltk.corpus import gutenberg
```

Output:
```
[nltk_data] Downloading package gutenberg to /root/nltk_data...
[nltk_data]   Package gutenberg is already up-to-date!
[nltk_data] Downloading package punkt_tab to /root/nltk_data...
[nltk_data]   Package punkt_tab is already up-to-date!
```

Complete the following function using the search function of the re library to find and return the year that a file in the gutenberg collection was written. Return the year as string or the empty string if not found.

<!-- grade_id: cell-6a452e90c99b682e, points: 0, grade: false, solution: true -->
```python
def get_year(file):
    text=gutenberg.raw(file)
    match=re.search(r'\b(1[5-9]\d{2}|20\d{2})\b', text)
    if match:
      return match.group(0)
    return""
```

<!-- grade_id: cell-7011f80019607313, points: 1, grade: true, solution: false -->
```python
"""Testing that the function returns the correct results"""
assert get_year('carroll-alice.txt') == '1865'
assert get_year('shakespeare-hamlet.txt') == '1599'
assert get_year('milton-paradise.txt') == '1667'
```

Complete the following function using the findall function of the re library to compute and report the number of times some text appears in a file of the gutenberg collection.

<!-- grade_id: cell-fcd8fa5f85a0dbd7, points: 0, grade: false, solution: true -->
```python
def find_mentions(text, file):
    periexomeno=gutenberg.raw(file)
    pattern=re.escape(text.strip())
    matches=re.findall(pattern, periexomeno)
    return len(matches)
```

<!-- grade_id: cell-b5c350fd93ad1d75, points: 1, grade: true, solution: false -->
```python
"""Testing that the function returns the correct results"""
assert find_mentions('CHAPTER ', 'carroll-alice.txt') == 12
assert find_mentions('UNK', 'shakespeare-caesar.txt') == 0
assert find_mentions('Caesar', 'shakespeare-caesar.txt') == 227
```

Complete the following function to return the mean number of words in the sentences of a gutenberg file that is given as input. The number should be rounded to 2 decimals.

<!-- grade_id: cell-4bc64ff6404ecf22, points: 0, grade: false, solution: true -->
```python
def average_sentence_length(file):
    sents = gutenberg.sents(file)
    lengths = [len(sent) for sent in sents]
    mean_length = sum(lengths) / len(lengths)
    return round(mean_length, 2)
    raise NotImplementedError()
```

<!-- grade_id: cell-27d42c9e86ec9169, points: 1, grade: true, solution: false -->
```python
"""Testing that the function returns the correct results"""
assert average_sentence_length('shakespeare-caesar.txt') == 11.94
assert average_sentence_length('austen-emma.txt') == 24.82
```

Output:
```
AssertionError                            Traceback (most recent call last)
/tmp/ipython-input-3551252387.py in <cell line: 0>()
      1 """Testing that the function returns the correct results"""
      2 assert average_sentence_length('shakespeare-caesar.txt') == 11.94
----> 3 assert average_sentence_length('austen-emma.txt') == 24.82

AssertionError:
```

## Part B: Byte-Pair Encoding (7 points)

Implement a function that returns the vocabulary learned by the byte-pair encoding algorithm given a corpus, after a number of iterations. The characters in the corpus, sorted in alphabetical order, will be the first members of the vocabulary, to be extended with as many tokens as the number of iterations given.

<!-- grade_id: cell-00f295323218bcab, points: 0, grade: false, solution: true -->
```python
from nltk import word_tokenize
def get_stats(tokens):
    pairs = {}
    for word, freq in tokens.items():
        symbols = word.split()
        for i in range(len(symbols)-1):
            pair = (symbols[i], symbols[i+1])
            pairs[pair] = pairs.get(pair, 0) + freq
    return pairs

def merge_pair(pair, tokens):
    new_tokens = {}
    bigram = ' '.join(pair)
    replacement = ''.join(pair)
    for word, freq in tokens.items():
        new_word = word.replace(bigram, replacement)
        new_tokens[new_word] = freq
    return new_tokens

def bpe_train(text, iterations):
    words = word_tokenize(text)
    tokens = {' '.join(['_'] + list(w)): text.split().count(w) for w in set(words)}

    vocab = sorted(set('_' + ''.join(words)))

    for _ in range(iterations):
        pairs = get_stats(tokens)
        if not pairs:
            break
        best = max(pairs, key=pairs.get)
        tokens = merge_pair(best, tokens)
        vocab.append(''.join(best))
    return vocab
        raise NotImplementedError()
```

<!-- grade_id: cell-9be88b50d1385e55, points: 7, grade: true, solution: false -->
```python
"""Testing that the function returns the correct results"""
corpus = "set new new renew reset renew"
assert set(bpe_train(corpus, 4)) == set(['_', 'e', 'n', 'r', 's', 't', 'w', 'ne', 'new', '_r', '_re'])
assert set(bpe_train(corpus, 8)) == set(['_', 'e', 'n', 'r', 's', 't', 'w', 'ne', 'new', '_r', '_re', 'se', 'set', '_new', '_renew'])
assert set(bpe_train(corpus, 9)) == set(['_', 'e', 'n', 'r', 's', 't', 'w', 'ne', 'new', '_r', '_re', 'se', 'set', '_new', '_renew', '_reset']);_
corpus = "hello world my large language model is coming after you beware"
assert set(bpe_train(corpus, 6)) == set(['_', 'a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'l', 'm', 'n', 'o', 'r', 's', 't', 'u', 'w', 'y', 'el', '_m', '_l', '_la', 'ge', 'ng'])
```